---
title       : generators
description : Generator functions
---
## Build a generator

```yaml
type: NormalExercise
lang: python
xp: 50
skills: 1
key: 93d17d2782
```

Generator functions are functions that yield a series of values, instead of returning a single value. A generator function uses the keyword yield instead of return.

In this exercise, you will create a generator function:

`@instructions`
- Complete the function header for the function get_name_lengths() that has a single parameter, names.
- In the for loop in the function definition, yield the length of the strings in names.
- Complete the iterable part of the for loop for printing the values generated by the get_name_lengths() generator function. Supply the call to get_name_lengths(), passing in the list python_guys.

`@hint`
- The syntax for a function header is def function name(parameters):.
- You can use the len() function to get the length of a string.
- Pass the list python_guys to a call to get_name_lengths(). Use this function call as the iterable for the for loop for printing values.

`@pre_exercise_code`
```{python}
python_guys = ['kamil', 'janusz', 'piotr', 'lukasz']
```

`@sample_code`
```{python}
# Create a list of strings
python_guys = ['kamil', 'janusz', 'piotr', 'lukasz']

# Define generator function get_name_lengths
def ____:
    """Generator function that yields the
    length of the strings in input_list."""

    # Yield the length of a string
    for person in names:
        ____

# Print the values generated by get_name_lengths()
for value in ____:
    print(value)

```

`@solution`
```{python}
# Create a list of strings
python_guys = ['kamil', 'janusz', 'piotr', 'lukasz']

# Define generator function get_name_lengths
def get_name_lengths(names):
    """Generator function that yields the
    length of the strings in input_list."""

    # Yield the length of a string
    for person in names:
        yield person

# Print the values generated by get_name_lengths()
for value in get_name_lengths(python_guys):
    print(value)

```

`@sct`
```{python}
test_object("python_guys",
            undefined_msg = "Make sure to define `python_guys`!",
            incorrect_msg = "Have you correctly assigned ``python_guys``!")
success_msg("Awesome! You correctly created your generator")
```
