---
title       : generators
description : Generator functions
---
## Build a generator

```yaml
type: NormalExercise
lang: python
xp: 50
skills: 1
key: 93d17d2782
```

Generator functions are functions that yield a series of values, instead of returning a single value. A generator function uses the keyword yield instead of return.

In this exercise, you will create a generator function.

`@instructions`
- Complete the function header with name `get_name_lengths` that has a single parameter `names`.
- yield the length of the strings in names inside loop
- Complete the iterable part of the for loop that includes the values generated by `get_name_lengths()`

`@hint`
- You can use the `len()` function to get the length of a string.
- Pass the list `cast_names` to a call `get_name_lengths()`.
- Use `get_name_lengths()` as the iterable for the for loop

`@pre_exercise_code`
```{python}
```

`@sample_code`
```{python}
# Define a list of strings
cast_names = ['Bolec', 'Grucha', 'Laska', 'Psikuta', 'Silnoreki']

# Define generator function get_name_lengths
def ____:
    """Generator function that yields the
    length of the strings in names list."""

    # Yield the length of a string
    for person in names:
        ____

# Print the values generated by get_name_lengths()
for value in ____:
    print(value)

```

`@solution`
```{python}
# Create a list of strings
cast_names = ['Bolec', 'Grucha', 'Laska', 'Psikuta', 'Silnoreki']

# Define generator function get_name_lengths
def get_name_lengths(names):
    """Generator function that yields the
    length of the strings in names list."""

    # Yield the length of a string
    for person in names:
        yield len(person)

# Print the values generated by get_name_lengths()
for value in get_name_lengths(cast_names):
    print(value)

```

`@sct`
```{python}
test_object("cast_names",
            undefined_msg = "Make sure to define `cast_names`!",
            incorrect_msg = "Don't modify the values of ``cast_names`` list!")
test_function("get_name_lengths",
            undefined_msg = "Definition of `get_name_lengths` function not created.",
            incorrect_msg = "Incorrt value of returned person length.")
#test_function("type(get_name_lengths(cast_names), generator)",
#			  incorrect_msg = "'get_name_lengths()' is not a generator")
success_msg("Awesome! You correctly created your generator")
```
